rule_files:
  - dashbase-alerts/api.yml

tests:
  - interval: 1m
    input_series:
      - series: jvm_cpu_usage_percent{component='api',app='test',instance='p1'}
        values: 41+0x10

    alert_rule_test:
      - eval_time: 10m
        alertname: api_cpu_high
        exp_alerts:
          - exp_labels:
              app: test
              instance: p1
              severity: warning
            exp_annotations:
              description: "test API instance p1 CPU usage high: 41%"

  - interval: 1m
    input_series:
      - series: jvm_memory_heap_usage{component='api',app='test',instance='p1'}
        values: 80+0x10

    alert_rule_test:
      - eval_time: 10m
        alertname: api_memory_high
        exp_alerts:
          - exp_labels:
              app: test
              instance: p1
              severity: warning
            exp_annotations:
              description: "test API instance p1 heap usage high: 80. Are there lots of queries?"

  - interval: 1m
    input_series:
      - series: dashbase_api_partition_latency_secs_count{app='test',root='foo',partition='1'}
        values: 1000+6000x30
      - series: dashbase_api_partition_timeouts_total{app='test',root='foo',partition='1'}
        values: 1000+360x30
      - series: dashbase_api_partition_latency_secs_count{app='test',root='foo',partition='2'}
        values: 1000+6000x30
      - series: dashbase_api_partition_timeouts_total{app='test',root='foo',partition='2'}
        values: 1000+299x30
    alert_rule_test:
      - eval_time: 31m
        alertname: api_partition_timeouts_high
        exp_alerts:
          - exp_labels:
              app: test
              root: foo
              partition: 1
              severity: critical
            exp_annotations:
              description: "test table foo partition 1 timeout fraction high: 0.06. Check the partition to see if it is under heavy load"

  - interval: 1m
    input_series:
      - series: dashbase_api_partition_latency_secs_count{app='test',root='foo',partition='1'}
        values: 1000+6000x30
      - series: dashbase_api_partition_errors_total{app='test',root='foo',partition='1'}
        values: 1000+360x30
      - series: dashbase_api_partition_latency_secs_count{app='test',root='foo',partition='2'}
        values: 1000+6000x30
      - series: dashbase_api_partition_errors_total{app='test',root='foo',partition='2'}
        values: 1000+299x30
    alert_rule_test:
      - eval_time: 31m
        alertname: api_partition_errors_high
        exp_alerts:
          - exp_labels:
              app: test
              root: foo
              partition: 1
              severity: critical
            exp_annotations:
              description: "test table foo partition 1 error fraction high: 0.06. Check the logs of the partition to see if there are lots of exceptions"

  - interval: 1m
    input_series:
      - series: dashbase_api_root_latency_secs_count{app='test',root='foo'}
        values: 1000+6000x30
      - series: dashbase_api_root_timeouts_total{app='test',root='foo'}
        values: 1000+360x30
      - series: dashbase_api_root_latency_secs_count{app='test',root='bar'}
        values: 1000+6000x30
      - series: dashbase_api_root_timeouts_total{app='test',root='bar'}
        values: 1000+299x30
    alert_rule_test:
      - eval_time: 31m
        alertname: api_root_timeouts_high
        exp_alerts:
          - exp_labels:
              app: test
              root: foo
              severity: critical
            exp_annotations:
              description: "test table foo root timeout fraction high: 0.06. Check the cpu/memory of the root."

  - interval: 1m
    input_series:
      - series: dashbase_api_root_latency_secs_count{app='test',root='foo'}
        values: 1000+6000x30
      - series: dashbase_api_root_errors_total{app='test',root='foo'}
        values: 1000+360x30
      - series: dashbase_api_root_latency_secs_count{app='test',root='bar'}
        values: 1000+6000x30
      - series: dashbase_api_root_errors_total{app='test',root='bar'}
        values: 1000+299x30
    alert_rule_test:
      - eval_time: 31m
        alertname: api_root_errors_high
        exp_alerts:
          - exp_labels:
              app: test
              root: foo
              severity: critical
            exp_annotations:
              description: "test table foo root error fraction high: 0.06. Check the logs of the root for exceptions"
