rule_files:
  - dashbase-alerts/kafka.yml

tests:
  - interval: 1m
    input_series:
      - series: jvm_memory_bytes_used{area='heap', component='kafka',app='staging',instance='kafka-0'}
        values: 990+0x5
      - series: jvm_memory_bytes_max{area='heap', component='kafka',app='staging',instance='kafka-0'}
        values: 1000+0x5

    alert_rule_test:
      - eval_time: 10m
        alertname: kafka_memory
        exp_alerts:
          - exp_labels:
              app: staging
              instance: kafka-0
              severity: warning
            exp_annotations:
              summary: "Kafka Memory Usage Too High In 'staging'"
              description: "Kafka instance 'kafka-0' in 'staging' is using +90% of available heap memory. Consider assigning more heap or scaling up Kafka."

  - interval: 1m
    input_series:
      - series: kube_persistentvolumeclaim_resource_requests_storage_bytes{app="kube-state-metrics",exported_namespace="monitor",instance="monitor-kube-state-metrics-8b4ddb64b-8zvt9",job="dashbase-pods",namespace="monitor",persistentvolumeclaim="kafka-data-kafka-0",pod_template_hash="460886206",release="monitor"}
        values: 1000+0x10
      - series: kafka_log_log_size{app="monitor",component="kafka",controller_revision_hash="kafka-6577d96494",instance="kafka-0",job="dashbase-pods",namespace="monitor",partition="0",statefulset_kubernetes_io_pod_name="kafka-0",topic="filebeat"}
        values: 990+1x10

    alert_rule_test:
      - eval_time: 10m
        alertname: kafka_disk
        exp_alerts:
          - exp_labels:
              app: monitor
              instance: kafka-0
              severity: critical
            exp_annotations:
              summary: "Kafka Running Out of Disk on 'kafka-0' in 'monitor'"
              description: "Kafka instance 'kafka-0' in 'monitor' is using +90% of available disk space. Consider reducing Kafka retention period or increase disk size."

  - interval: 1m
    input_series:
      - series: kafka_server_brokertopicmetrics_bytesrejected_total{component='kafka',app='staging',instance='kafka-0',topic='abc'}
        values: 0+10x10

    alert_rule_test:
      - eval_time: 10m
        alertname: kafka_bytes_rejected
        exp_alerts:
          - exp_labels:
              app: staging
              topic: abc
              severity: critical
            exp_annotations:
              summary: "Kafka Too Many Bytes Rejected In 'staging'"
              description: "Kafka is rejecting bytes for topic 'abc' in 'staging'. Check the status of the producer and broker."
