rule_files:
  - dashbase-alerts/proxy.yml

tests:
  - interval: 1m
    input_series:
      - series: jvm_cpu_usage_percent{component='proxy',app='test'}
        values: 90+0x5

    alert_rule_test:
      - eval_time: 10m
        alertname: proxy_cpu
        exp_alerts:
          - exp_labels:
              app: test
              severity: warning
            exp_annotations:
              summary: "Proxy CPU Usage Too High in 'test'"
              description: "Average Proxy CPU usage in 'test' is over 75%. Consider scaling up Proxy."

  - interval: 1m
    input_series:
      - series: jvm_memory_heap_usage{component='proxy',app='test',instance='p1'}
        values: 0.9+0x5
      - series: jvm_memory_heap_usage{component='proxy',app='test',instance='p2'}
        values: 0.99+0x5

    alert_rule_test:
      - eval_time: 10m
        alertname: proxy_memory
        exp_alerts:
          - exp_labels:
              app: test
              severity: warning
            exp_annotations:
              summary: "Proxy Memory Usage Too High in 'test'"
              description: "Average Proxy Memory usage in 'test' is over 80%. Consider scaling up Proxy."

  - interval: 1m
    input_series:
      - series: dashbase_proxy_events_parse_error{component='proxy',app='test',topic='filebeat'}
        values: 0+1x10
      - series: dashbase_proxy_events_received{component='proxy',app='test',topic='filebeat'}
        values: 0+100x10

    alert_rule_test:
      - eval_time: 10m
        alertname: proxy_parse_error
        exp_alerts:
          - exp_labels:
              app: test
              topic: filebeat
              severity: critical
            exp_annotations:
              summary: "Proxy Too Many Parse Errors for topic 'filebeat'"
              description: "Proxy failed to parse more than 1% of the incoming messages for topic 'filebeat' in 'test'. Check Parser's Config."

  - interval: 1m
    input_series:
      - series: dashbase_proxy_event_delay_bucket{component='proxy',app='test',topic='filebeat', le="3600.0"}
        values: 0+0x10
      - series: dashbase_proxy_event_delay_count{component='proxy',app='test',topic='filebeat'}
        values: 0+1000x10

    alert_rule_test:
      - eval_time: 10m
        alertname: proxy_too_many_old_message
        exp_alerts:
          - exp_labels:
              app: test
              topic: filebeat
              severity: warning
            exp_annotations:
              summary: "Proxy Too Many Old Messages for topic 'filebeat'"
              description: "More than 50% of messages received by Proxy are more than 1 hour old for topic 'filebeat' in 'test'. Check the status of message producer."

  - interval: 1m
    input_series:
      - series: dashbase_proxy_event_size_bucket{component='proxy',app='test',topic='filebeat', le="1.048576E7"}
        values: 0+0x10
      - series: dashbase_proxy_event_size_count{component='proxy',app='test',topic='filebeat'}
        values: 0+1000x10

    alert_rule_test:
      - eval_time: 10m
        alertname: proxy_too_many_large_message
        exp_alerts:
          - exp_labels:
              app: test
              topic: filebeat
              severity: warning
            exp_annotations:
              summary: "Proxy Too Large Messages for topic 'filebeat'"
              description: "More than 1% of messages received by Proxy are bigger than 10MB for topic 'filebeat' in 'test'. Check the configuration of message producer."

  - interval: 1m
    input_series:
      - series: dashbase_kafka_failure{component='proxy',app='test',exception='NullPointerException'}
        values: 0+10x10
      - series: dashbase_kafka_failure{component='proxy',app='test',exception='IllegalArgumentException'}
        values: 0+10x10
      - series: dashbase_kafka_success{component='proxy',app='test',topic='filebeat'}
        values: 0+100x10

    alert_rule_test:
      - eval_time: 10m
        alertname: proxy_kafka_error
        exp_alerts:
          - exp_labels:
              app: test
              exception: NullPointerException
              severity: critical
            exp_annotations:
              summary: "Proxy Too Many Kafka Errors"
              description: "Proxy are seeing too many errors sending messages to kafka with exception 'NullPointerException' in 'test'. Check the status of kafka and also Proxy's configuration for Kafka."
          - exp_labels:
              app: test
              exception: IllegalArgumentException
              severity: critical
            exp_annotations:
              summary: "Proxy Too Many Kafka Errors"
              description: "Proxy are seeing too many errors sending messages to kafka with exception 'IllegalArgumentException' in 'test'. Check the status of kafka and also Proxy's configuration for Kafka."

