rule_files:
  - dashbase-alerts/zookeeper.yml

tests:
  - interval: 1m
    input_series:
      - series: jvm_memory_bytes_used{area='heap', component='zookeeper',app='staging',instance='kafka-0'}
        values: 990+0x5
      - series: jvm_memory_bytes_max{area='heap', component='zookeeper',app='staging',instance='kafka-0'}
        values: 1000+0x5

    alert_rule_test:
      - eval_time: 10m
        alertname: zookeeper_memory
        exp_alerts:
          - exp_labels:
              app: staging
              severity: warning
            exp_annotations:
              summary: "ZooKeeper Memory Usage Too High In 'staging'"
              description: "ZooKeeper in 'staging' is using +90% of available heap memory. Consider assigning more heap."

  - interval: 1m
    input_series:
      - series: zookeeper_OutstandingRequests{app='staging'}
        values: 0+10x10

    alert_rule_test:
      - eval_time: 10m
        alertname: zookeeper_outstanding_requests
        exp_alerts:
          - exp_labels:
              app: staging
              severity: warning
            exp_annotations:
              summary: "ZooKeeper Outstanding Requests In 'staging'"
              description: "There are too many requests to ZooKeeper in 'staging'. Check the status of Dashbase API and Table."
